# Lab-1 - Some Basic Growth Facts

## Purpose

This introductory lab will serve two purposes.  

1. It will acquaint you with some basic growth facts that you will seek to explain later in the semester.  
2. It will allow you to apply your new knowledge of Python and Jupyter notebooks.

# Assignment

In this lab, you will document cross-country differences in income and growth rates to help motivate the study of growth throughout the semester. The final product is a Jupyter notebook that provides and answer to all the questions. It must perform all analyses and create all the figures, and export the outputs requested. 

# Setup -  Import packages and setup paths

Create a `pathout` variable and `pathgraphs` for the name of the directories where you will save data and figures. If they do not exist, make sure to create them

<!-- grade_id: cell-1f18136b34393c80, points: 0, grade: false, solution: true -->
```python
# Let's import pandas, os, numpy, and some other basic packages we will use 
# YOUR CODE HERE
raise NotImplementedError()

# Create output directories
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b1aa21e6c45d7cea, points: 2.5, grade: true, solution: false -->
```python
assert(os.path.exists(pathout)==True)
assert(os.path.exists(pathgraphs)==True)
```

# 1. Import Data

Use the code from class to create income per capita data from the Penn World Tables and the World Development Indicators for the 1950-2021 period.

Create a dataframe for each data source that contains the following variables:
1. Country code: `country_code`
2. Country name: `country_name`
2. Year: `year`
3. GDP per capita: `gdp_pc`

For PWT use only `rgdpe` for GDP and name the dataframe `pwt`, for WDi use `'NY.GDP.PCAP.PP.KD'` for GDP per capita and name it `wdi`.

Save the files with names `pwt100.dta` and `wdi.dta` in the `pathout` directory.

## PWT-10

<!-- grade_id: cell-106b8ca8b19737cb, points: 0, grade: false, solution: true -->
```python
# Download PWT-10 data
# YOUR CODE HERE
raise NotImplementedError()

# Generate GDPpc using rgdpe measure
# YOUR CODE HERE
raise NotImplementedError()

# Select columns and rename them if necessary
# YOUR CODE HERE
raise NotImplementedError()

# Save dataframe with 4 requested variables
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9be909839ff1dbde, points: 5, grade: true, solution: false -->
```python
assert(pwt.shape==(12810,4))
assert(os.path.exists(pathout + '/pwt100.dta')==True)
```

## WDI

Get countries and their regions and basic info useful later, especially for plotting

<!-- grade_id: cell-32c051ffcc13b92d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Download data

<!-- grade_id: cell-a8f3321fbe617376, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Merge WDI datasets so we have codes and names together

<!-- grade_id: cell-87ee6c45fa2f04e5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Select the 4 requested variables and save file

<!-- grade_id: cell-5993e1eaa1dea3d6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ed4011b449d77532, points: 5, grade: true, solution: false -->
```python
assert(wdi.shape==(13454,4))
assert(os.path.exists(pathout + '/wdi.dta')==True)
assert((wdi.columns==['country_code', 'country_name', 'year', 'gdp_pc']).all())
```

# 2. Compare Income levels

For each of these two datasets: 

* compute the average level of real per capita income in the earliest and in the latest year across all the countries in each dataset.  
* Find the minimum per capita income in each of these two periods
* Compute the maximum per capita income in each of these two periods
* Compute the ratio of GDP per capita between richest to poorest in the two periods
* Show how the distribution of income has changed over time by  plotting the distributions of log-income per capita in different years using kernel densities. Export each figure as a pdf file into ``pathgraphs`` with name ``dataset_gdppc_year.pdf``. To do this, write a function that takes as inputs the dataset and the year you want to plot.
* Create a figure that compares the distribution of log-income per capita in both years for each dataset, i.e., it shows both jointly in one figure, and export it as a pdf and a png file into ``pathgraphs`` with names ``dataset_gdppc_minyear_maxyear.pdf`` and ``dataset_gdppc_minyear_maxyear.png``.

You need to name the variables in a way that the example code runs. 

Save the earliest year in PWT as `year_min_pwt`.

<!-- grade_id: cell-afc2623d777d61a3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f5ddd6c5017b35c7, points: 1, grade: true, solution: false -->
```python
print("Earliest year in PWT is ", year_min_pwt, '\n')
```

Earliest year in WDI as `year_min_wdi`

<!-- grade_id: cell-4b1a8c46621f3029, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e567b309b46cab0d, points: 1, grade: true, solution: false -->
```python
print("Earliest year in WDI is ", year_min_wdi, '\n')
```

Latest year in PWT as `year_max_pwt`

<!-- grade_id: cell-13da7bcc2b029d4b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c87de0e2fbfc491c, points: 1, grade: true, solution: false -->
```python
print("Latest year in PWT is ", year_max_pwt, '\n')
```

Latest year in WDI as `year_max_wdi`

<!-- grade_id: cell-0f3090c4ac03a294, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ffbf9548f485bccd, points: 1, grade: true, solution: false -->
```python
print("Latest year in WDI is ", year_max_wdi, '\n')
```

Average GDP pc in earliest year in PWT as `mean_gdppc_pwt_early`

<!-- grade_id: cell-5e01577356a7efd9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2e11fc7a01cc7420, points: 1, grade: true, solution: false -->
```python
print("The average income per capita in the earliest year, i.e.,", year_min_pwt, "in the PWT, is ", mean_gdppc_pwt_early, '\n')
```

Average GDP pc in earliest year in WDI as `mean_gdppc_wdi_early`

<!-- grade_id: cell-a008aac4b29d6ed8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-62297975f3f5bf56, points: 1, grade: true, solution: false -->
```python
print("The average income per capita in the earliest year, i.e.,", year_min_wdi, "in the WDI, is ", mean_gdppc_wdi_early, '\n')
```

Average GDP pc in latest year in PWT as `mean_gdppc_pwt_late`

<!-- grade_id: cell-793af247164dac85, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c271b86663463a1e, points: 1, grade: true, solution: false -->
```python
print("The average income per capita in the latest year, i.e.,", year_max_pwt, "in the PWT, is ", mean_gdppc_pwt_late, '\n')
```

Average GDP pc in latest year in WDI as `mean_gdppc_wdi_late`

<!-- grade_id: cell-865d23fd2e756892, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-23fce009562d37a6, points: 1, grade: true, solution: false -->
```python
print("The average income per capita in the latest year, i.e.,", year_max_wdi, "in the WDI, is ", mean_gdppc_wdi_late, '\n')
```

Minimum GDP pc in earliest year in PWT as `min_gdppc_pwt_early`

<!-- grade_id: cell-0e4cd8962831097b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c203c74cc486219b, points: 1, grade: true, solution: false -->
```python
print("The minimum income per capita in the earliest year, i.e.,", year_min_pwt, "in the PWT, is ", min_gdppc_pwt_early, '\n')
```

Minimum GDP pc in earliest year in WDI as `min_gdppc_wdi_early`

<!-- grade_id: cell-c1ececdca468482c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-818e66b92095b995, points: 1, grade: true, solution: false -->
```python
print("The minimum income per capita in the earliest year, i.e.,", year_min_wdi, "in the WDI, is ", min_gdppc_wdi_early, '\n')
```

Minimum GDP pc in latest year in PWT as `min_gdppc_pwt_late`

<!-- grade_id: cell-ef74288b9b6dbe8c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0cf9e07785a0b5a5, points: 1, grade: true, solution: false -->
```python
print("The minimum income per capita in the latest year, i.e.,", year_max_pwt, "in the PWT, is ", min_gdppc_pwt_late, '\n')
```

Minimum GDP pc in latest year in WDI as `min_gdppc_wdi_late`

<!-- grade_id: cell-41938ef77dde4117, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b7fd8dc0bab0a6bb, points: 1, grade: true, solution: false -->
```python
print("The minimum income per capita in the latest year, i.e.,", year_max_wdi, "in the WDI, is ", min_gdppc_wdi_late, '\n')
```

Maximum GDP pc in earliest year in PWT as `max_gdppc_pwt_early`

<!-- grade_id: cell-87fbacd7d7fad503, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d0d9f60799628d67, points: 1, grade: true, solution: false -->
```python
print("The maximum income per capita in the earliest year, i.e.,", year_min_pwt, "in the PWT, is ", max_gdppc_pwt_early, '\n')
```

Maximum GDP pc in earliest year in WDI as `max_gdppc_wdi_early`

<!-- grade_id: cell-0c5a8189069a91bf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ade2c564e718253e, points: 1, grade: true, solution: false -->
```python
print("The maximum income per capita in the earliest year, i.e.,", year_min_wdi, "in the WDI, is ", max_gdppc_wdi_early, '\n')
```

Maximum GDP pc in latest year in PWT as `max_gdppc_pwt_late`

<!-- grade_id: cell-d3988950ccb52994, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-33296a0e0170e321, points: 1, grade: true, solution: false -->
```python
print("The maximum income per capita in the latest year, i.e.,", year_max_pwt, "in the PWT, is ", max_gdppc_pwt_late, '\n')
```

Maximum GDP pc in latest year in WDI as `max_gdppc_wdi_late`

<!-- grade_id: cell-73fa3e093dd67a94, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2246ac42b3c5d740, points: 1, grade: true, solution: false -->
```python
print("The maximum income per capita in the latest year, i.e.,", year_max_wdi, "in the WDI, is ", max_gdppc_wdi_late, '\n')
```

Ratio of GDP pc in earliest year in PWT as `ratio_gdppc_pwt_early`

<!-- grade_id: cell-05c62cd437c0eb72, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fa9d76c13546ff92, points: 1, grade: true, solution: false -->
```python
print("The ratio income per capita in the earliest year, i.e.,", year_min_pwt, "in the PWT, is ", ratio_gdppc_pwt_early, '\n')
```

Ratio of GDP pc in earliest year in WDI as `ratio_gdppc_wdi_early`

<!-- grade_id: cell-bccf344ade5fdad4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c3d2fc191d17efeb, points: 1, grade: true, solution: false -->
```python
print("The ratio income per capita in the earliest year, i.e.,", year_min_wdi, "in the WDI, is ", ratio_gdppc_wdi_early, '\n')
```

Ratio of GDP pc in latest year in PWT as `ratio_gdppc_pwt_late`

<!-- grade_id: cell-ba9fa189cc5558c7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4d114926f1ca9a5c, points: 1, grade: true, solution: false -->
```python
print("The ratio income per capita in the latest year, i.e.,", year_max_pwt, "in the PWT, is ", ratio_gdppc_pwt_late, '\n')
```

Ratio of GDP pc in latest year in WDI as `ratio_gdppc_wdi_late`

<!-- grade_id: cell-0ad7da478f2c215c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4c4e2928a6ee7bb5, points: 1, grade: true, solution: false -->
```python
print("The ratio income per capita in the latest year, i.e.,", year_max_wdi, "in the WDI, is ", ratio_gdppc_wdi_late, '\n')
```

<!-- grade_id: cell-5a54b444406bec1b, points: 5, grade: true, solution: false -->
```python
assert(year_min_pwt==1950)
assert(year_max_pwt==2019)

```

Let's now create the function to plot each distribution in each year and dataset

<!-- grade_id: cell-aa7cb53e6afe319f, points: 0, grade: false, solution: true -->
```python
sns.set(rc={'figure.figsize':(11.7,8.27)})
#sns.reset_orig()
sns.set_context("talk")

def my_kde_plots(dataset='pwt', myyear=year_min_pwt):
    '''
    This function plots the KDE-plot of log income per capita, where
    
    dataset: is either pwt or wdi (one of the two datasets you created above)
    myyear: is one of the year variables you created above, e.g., year_min_pwt, year_min_wdi
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6e064c037304d5a8, points: 5, grade: true, solution: false -->
```python
# KDE plot for PWT in earliest period
my_kde_plots()
```

<!-- grade_id: cell-daf7e60815921e4b, points: 5, grade: true, solution: false -->
```python
# KDE plot for PWT in latest period
my_kde_plots(myyear=year_max_pwt)
```

<!-- grade_id: cell-2244a0c19379d0a7, points: 5, grade: true, solution: false -->
```python
# KDE plot for WDI in earliest period
my_kde_plots(dataset='wdi', myyear=year_min_wdi)
```

<!-- grade_id: cell-bf4af7a308022573, points: 5, grade: true, solution: false -->
```python
# KDE plot for PWT in latest period
my_kde_plots(dataset='wdi', myyear=year_max_wdi)
```

<!-- grade_id: cell-9e89765762b26403, points: 2.5, grade: true, solution: false -->
```python
assert(os.path.exists(pathgraphs + '/pwt_gdppc_' + str(year_min_pwt) + '.pdf')==True)
assert(os.path.exists(pathgraphs + '/wdi_gdppc_' + str(year_min_wdi) + '.pdf')==True)
```

Let's write a function that plots for each dataset and pair of years the distributions in that year and exports it to both png and pdf.

<!-- grade_id: cell-20032e775009fa97, points: 0, grade: false, solution: true -->
```python
sns.set(rc={'figure.figsize':(11.7,8.27)})
#sns.reset_orig()
sns.set_context("talk")

def my_kde_plots_joint(dataset='pwt', myyears=[year_min_pwt, year_max_pwt]):
    '''
    This function plots the KDE-plot of log income per capita in two years, where
    
    dataset: is either pwt or wdi (one of the two datasets you created above)
    myyears: is a list of years
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f93b853f045b331d, points: 2.5, grade: true, solution: false -->
```python
# Plot PWT
my_kde_plots_joint()
```

<!-- grade_id: cell-d7ceffb95f6560bf, points: 2.5, grade: true, solution: false -->
```python
# Plot WDI
my_kde_plots_joint(dataset='wdi', myyears=[year_min_wdi, year_max_wdi])
```

<!-- grade_id: cell-ef075b38ad1899b7, points: 2.5, grade: true, solution: false -->
```python
assert(os.path.exists(pathgraphs + 'pwt_gdppc_' + str(year_min_pwt) + '_' + str(year_max_pwt) + '.pdf')==True)
assert(os.path.exists(pathgraphs + 'wdi_gdppc_' + str(year_min_wdi) + '_' + str(year_max_wdi) + '.pdf')==True)
```

# 3. Growth rates

Compute the growth rate of real per capita income over the time period between the first year and the last year as

$$
growth = \log(gdp\_pc_{last}) - \log(gdp\_pc_{first})
$$

Save the results in a new dataframe called `growth_dataset`.  What is the average growth rate for all the countries in your data? Save it as `mean_growth_dataset`. What is the minimum growth rate? Save it as `min_growth_dataset`. What is the maximum growth rate? Save it as `max_growth_dataset`.

Compute growth rates

<!-- grade_id: cell-d1c97f59295e8c0b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Compute Average growth rate

<!-- grade_id: cell-549f20d20cf21923, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f14f47186f4c69f6, points: 2.5, grade: true, solution: false -->
```python
print('Average growth rate between', year_min_pwt, 'and', year_max_pwt, 'according to PWT was {0:.2%}'.format(mean_growth_pwt))
print('Average growth rate between', year_min_wdi, 'and', year_max_wdi, 'according to WDI was {0:.2%}'.format(mean_growth_wdi))

```

Compute minimum growth rate

<!-- grade_id: cell-d56ff2ee06d45c3e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-55f6b4e018482c00, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Compute maximum growth rate

<!-- grade_id: cell-d5b1c547fe35cb21, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0f079ace95a009a2, points: 2.5, grade: true, solution: false -->
```python
print('Maximum growth rate between', year_min_pwt, 'and', year_max_pwt, 'according to PWT was {0:.2%}'.format(max_growth_pwt))
print('Maximum growth rate between', year_min_wdi, 'and', year_max_wdi, 'according to WDI was {0:.2%}'.format(max_growth_wdi))

```

# 4. Growth rates by WB region

Consider the regions defined by the World Bank: East Asia & Pacific, Europe & Central Asia, Latin America & Caribbean, Middle East & North Africa, North America, South Asia, and Sub-Saharan Africa.  Compute the minimum, maximum, and mean growth rates of income per capita for countries in each of these groups in the earliest and latest year in each dataset.  What patterns can you identify in the growth of income and in the dispersion of income in these regions?

Start by merging the `wbcountries` dataframe with the others to get the regional variable. Create new dataframes `growth_dataset_regions`. 

**Note:** There are 2 countries in PWT that do not appear in the `wbcountries` dataframe. 
They both belong to Latin America & Caribbean (codes: `AIA`, `MSR`). So, they will be dropped in the `joint` merge.

<!-- grade_id: cell-633e45e6582f63b8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-21ad096074fb87b3, points: 2.5, grade: true, solution: false -->
```python
assert(growth_wdi.shape[0]==growth_wdi_regions.shape[0])
assert((growth_pwt.shape[0] - 2)==growth_pwt_regions.shape[0])
```

Compute the statistics by region and save them in a new dataframe `statistic_growth_dataset_regions` that has 1 column with the regions and another with the growth rates.

<!-- grade_id: cell-0bd781c17aa480fa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# 5. Counter-factual Analysis

Using the PWT, take the maximum growth rate that you computed in question 3 and apply it to the poorest country in the earliest date.  Then take the minimum growth rate that you computed in question 3 and apply it to the richest country in the earliest date.  How would these different growth rates affect the level of income in these countries in the last period?  Explain how this computation demonstrates the importance of the study of long-run growth.

Compute counter factual income per capita for the poorest country. Safe value in `poor_counter_factual`.

<!-- grade_id: cell-e4034c55f1e16621, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Compute counter factual income per capita for the richest country. Safe value in `rich_counter_factual`.

<!-- grade_id: cell-937c935aa8cba799, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-55559258d246fec5, points: 5, grade: true, solution: false -->
```python
assert(rich_counter_factual<poor_counter_factual)
```

# 6. Who's who in economic growth?

Sort the data by the growth rates you computed in number 3.  Identify the 10 fastest growing and the 10 slowest growing countries.  Do the fast growers have anything in common?  Do the slow growers?  Do these lists suggest potential ideas for investigation of the determinants of growth?

Let's sort and show the top 10 slowest growing countries in PWT

<!-- grade_id: cell-b98c4ebf00f67a32, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Now in WDI

<!-- grade_id: cell-6fd548697e9f0ec1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Let's sort and show the top 10 fastest growing countries in PWT

<!-- grade_id: cell-7881cdc32535f82e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Now in WDI

<!-- grade_id: cell-5413cde488f968b7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# 7. Convergence

Generate a `plotly` scatter plot that shows the income level in the earliest period on the horizontal axis and the growth rate between earliest and latest in the vertical axis. Make sure hovering on a symbols shows the information of the country. Use colors for WB regions and shapes for WB income groups. Export it to file `dataset_convergence_early_late.pdf`. Do you see anything noteworthy?

Start by creating a dataframe with only data for the earliest date `dataset_early` and then merge it with the growth rates dataframe `growth_dataset` (make sure to add a suffix to distinguish variables if needed). Drop any NA's to simplify plotting.

First PWT

<!-- grade_id: cell-4ea1badb9248e4f2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Now WDI

```python
wdi_early = wdi.loc[wdi.year==year_min_wdi].reset_index(drop=True).copy()
wdi_early = wdi_early.merge(growth_wdi, on=['country_code', 'country_name'], suffixes=['', '_growth'])
wdi_early = wdi_early.merge(wbcountries[['iso3c', 'region', 'incomeLevel']], left_on='country_code', right_on='iso3c')
wdi_early = wdi_early.dropna().reset_index(drop=True)
wdi_early
```

Now let's plot this data for PWT

<!-- grade_id: cell-00b095ab7125e6b5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0e6126c984bab237, points: 5, grade: true, solution: false -->
```python
assert(os.path.exists(pathgraphs + 'pwt_convergence_' + str(year_min_pwt) + '_' + str(year_max_pwt) + '.pdf'))
```

Now for WDI

<!-- grade_id: cell-c91a2028c33611e6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cbf1b575e9cddb9d, points: 5, grade: true, solution: false -->
```python
assert(os.path.exists(pathgraphs + 'wdi_convergence_' + str(year_min_wdi) + '_' + str(year_max_wdi) + '.pdf'))
```

# 8. Persistence

Generate a `plotly` scatter plot that shows the income level in the earliest period on the horizontal axis and the income level in the latest in the vertical axis. Make sure hovering on a symbols shows the information of the country. Use colors for WB regions and shapes for WB income groups. Export it to file `dataset_persistence_early_late.pdf`. Do you see anything noteworthy?

Start by creating a dataframe with only data for the latest date `dataset_late` and then merge it with the `dataset_early` dataframe you cerated in teh previous point (make sure to add a suffix to distinguish variables if needed). Drop any NA's to simplify plotting.

Start with PWT

<!-- grade_id: cell-187179cae74cdcbd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Now WDI

<!-- grade_id: cell-f88aac24a08de9b4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Now let's plot this data for PWT

<!-- grade_id: cell-90b1bf6efdd50558, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3434c9dc87ec5a10, points: 5, grade: true, solution: false -->
```python
assert(os.path.exists(pathgraphs + 'pwt_persistence_' + str(year_min_pwt) + '_' + str(year_max_pwt) + '.pdf'))
```

Now WDI

<!-- grade_id: cell-0f954ffbfb90b55a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7ffcfe5a3ccd96ef, points: 5, grade: true, solution: false -->
```python
assert(os.path.exists(pathgraphs + 'wdi_persistence_' + str(year_min_pwt) + '_' + str(year_max_pwt) + '.pdf'))
```